### Q1
>- Import the packages `os`, and `pandas`

<!-- grade_id: cell-b7831d6fbf7d613d, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import os
```

<!-- grade_id: cell-45779676dba5b018, points: 0, grade: true, solution: false -->
```python

```

### Q2
>- Write a quick loop using `os` to find the file `spotify.csv`
>- store the file name in a variable called `file`

<!-- grade_id: cell-70741dad6ab3471a, points: 0, grade: false, solution: true -->
```python
os.getcwd()
for file in os.listdir():
    if '.csv' in file:
        print(file)
    elif '.xlsx' in file:
        print(file)
```

Output:
```
spotify_tracks.csv
```

<!-- grade_id: cell-e479b00815329dd6, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
assert file == 'spotify_tracks.csv'
        
### END TESTS
```

### Q3

>- Read the file `spotify_tracks.csv` into python as a pandas dataframe.
>- Make sure that you remove the index column using `index_col = 0`
>- Name your dataframe `spotify`

<!-- grade_id: cell-5b5a7269f2cc4787, points: 0, grade: false, solution: true -->
```python
spotify = pd.read_csv(file, index_col = 0)
spotify.head()
```

Output:
```
acousticness                album_id  \
0         0.294  0D3QufeCudpQANOR7luqdr   
1         0.863  1bcqsH5UyTBzmh9YizdsBE   
2         0.750  4tKijjmxGClg4JOLAyo2qE   
3         0.763  6FeJF5r8roonnKraJxr4oB   
4         0.770  4tKijjmxGClg4JOLAyo2qE   

                                        analysis_url  \
0  https://api.spotify.com/v1/audio-analysis/5qlj...   
1  https://api.spotify.com/v1/audio-analysis/3VAX...   
2  https://api.spotify.com/v1/audio-analysis/1L3Y...   
3  https://api.spotify.com/v1/audio-analysis/6aCe...   
4  https://api.spotify.com/v1/audio-analysis/1Vo8...   

                   artists_id  \
0  ['3mxJuHRn2ZWD5OofvJtDZY']   
1  ['4xWMewm6CYMstu0sPgd9jJ']   
2  ['3hYaK5FF3YAglCj5HZgBnP']   
3  ['2KQsUB9DRBcJk17JWX1eXD']   
4  ['3hYaK5FF3YAglCj5HZgBnP']   

                                   available_markets country  danceability  \
0  ['AD', 'AE', 'AR', 'AT', 'AU', 'BE', 'BG', 'BH...      BE         0.698   
1  ['AD', 'AE', 'AR', 'AT', 'AU', 'BE', 'BG', 'BH...
```

### DO NOT RUN THE NEXT CELL

```python
################################
################################
##### DO NOT RUN THIS CELL #####
################################
################################

### This cell is so my computer will load the proper data when I grade the assignment
### If you run this cell, simply re run the cell you used to read in the data (it should be the cell above this one)

spotify = pd.read_csv('/Users/drewwilliams/Desktop/PyClass/2021 Fall/spotify_tracks.csv', index_col = 0)
```

### CONTINUE FROM HERE

<!-- grade_id: cell-1a1618a6a3f87509, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
assert spotify.shape == (101939, 31)
        
### END TESTS
```

### Q4
>- Use `.shape` to tell me how many rows and columns are in the dataframe
>- Store the number of rows in a variable called `rows`
>- Store the number of columns in a variable called `cols`
>- Look at the first 3 rows of data

<!-- grade_id: cell-077cb637072c1ae2, points: 0, grade: false, solution: true -->
```python
print(dat.shape)
spotify.head(3)

rows = spotify.shape[0]
columns =spotify.shape[1]


print(rows, columns)
```

Output:
```
(101939, 31)
101939 31
```

<!-- grade_id: cell-3f382313464de68d, points: 2, grade: true, solution: false -->
```python

```

### Q5
>- How many null values are in each column?

<!-- grade_id: cell-c24e4fd80bbd175e, points: 0, grade: false, solution: true -->
```python
spotify.isnull().sum()

```

Output:
```
acousticness         0
album_id             0
analysis_url         0
artists_id           0
available_markets    0
country              0
danceability         0
disc_number          0
duration_ms          0
energy               0
href                 0
id                   0
instrumentalness     0
key                  0
liveness             0
loudness             0
lyrics               0
mode                 0
name                 0
playlist             0
popularity           0
preview_url          0
speechiness          0
tempo                0
time_signature       0
track_href           0
track_name_prev      0
track_number         0
uri                  0
valence              0
type                 0
dtype: int64
```

<!-- grade_id: cell-929c34bb92744e6c, points: 3, grade: true, solution: false -->
```python

```

### Q6
>- Get the descriptive analytics for the entire dataframe
>- Store the mean of `liveness` in a variable called `mean_liveness`
>- HINT: Do this by filtering the dataframe, not by copy pasting the value from it

<!-- grade_id: cell-542c39500e9a22c9, points: 0, grade: false, solution: true -->
```python
#print(dat.describe())
mean_liveness = dat.describe().loc['mean']['liveness']
mean_liveness
```

Output:
```
0.19764030645778108
```

<!-- grade_id: cell-b2385ea487464de8, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
assert mean_liveness == spotify.describe().loc['mean']['liveness']
### END HIDDEN TEST
```

### Q7
>- How many unique artists are there? Store the result in the variable `unique_artists`
>- How many unique albums are there? Store the result in the variable `unique_albums`

<!-- grade_id: cell-ad404d52ccd4b2cb, points: 0, grade: false, solution: true -->
```python
unique_artists = spotify['artists_id'].nunique()
unique_albums = spotify['album_id'].nunique()

print(unique_artists)
unique_albums
```

Output:
```
54673
75503
```

<!-- grade_id: cell-a340ea16d763f9e8, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
assert unique_artists == 54673
assert unique_albums == 75503
### END HIDDEN TEST
```

### Q8
>- What are the top 5 artists IDs based on count?
>- Store this series in a variable called `top_5_artists`
>- Make sure your series is ordered in Descending Order. Only include the top 5 `artists_ids`

<!-- grade_id: cell-47bdd6122f44ca73, points: 0, grade: false, solution: true -->
```python
top_5_artists = spotify['artists_id'].value_counts().sort_values(ascending = False)[0:5]

top_5_artists
```

Output:
```
['0LyfQWJT6nXafLPZqxe9Of']    284
['1l6d0RIxTL3JytlLGvWzYe']    251
['0vLsqW05dyLvjuKKftAEGA']    180
['49FYG2DSZqIVpCtoir2Upw']    162
['3iiABQlNuezpR3DcrlyPoA']    158
Name: artists_id, dtype: int64
```

<!-- grade_id: cell-a7bb41c85526c255, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
assert top_5_artists[0] == 284
assert top_5_artists[-1] == 158
### END HIDDEN TEST
```

### Q9
>- Create a dataframe that only has data from the country `AR`
>- Name your new dataframe `spotify_ar`
>- Store the number of columns in a variable called `n_cols`
>- Store the number of rows in a variable called `n_rows`

<!-- grade_id: cell-dffb6e1b3f237be8, points: 0, grade: false, solution: true -->
```python
spotify_ar = pd.DataFrame(spotify[spotify['country'] == 'AR'])

print(spotify_ar.shape)

n_cols = spotify_ar.shape[1]
n_rows = spotify_ar.shape[0]





```

Output:
```
(50437, 31)
50437
```

<!-- grade_id: cell-66709d13dbbb54bc, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
assert n_rows == 50437
assert n_cols == 31
### END HIDDEN TEST
```

### Q10
>- Use `spotify` (not `spotify_ar`)
>- What is the count, average, median, max, and minimum duration based on track number?
>- HINT: use `.agg()`
>- Store this new dataframe in a variable called `track_stats`

<!-- grade_id: cell-898472a6b8a6721d, points: 0, grade: false, solution: true -->
```python
track_stats = spotify.groupby('track_number').duration_ms.agg(['count','mean','median','max','min'])
track_stats
```

Output:
```
count           mean    median        max       min
track_number                                                     
1.0           45300  237628.788830  211086.5  5505831.0    4000.0
2.0           10559  267595.584620  227480.0  4057967.0    5400.0
3.0            7715  266678.384057  228240.0  4342462.0    3500.0
4.0            6057  273238.258049  229436.0  4373004.0    5375.0
5.0            4980  260390.310241  228720.0  3614407.0    4250.0
...             ...            ...       ...        ...       ...
231.0             1  196680.000000  196680.0   196680.0  196680.0
244.0             1  146680.000000  146680.0   146680.0  146680.0
327.0             1   17244.000000   17244.0    17244.0   17244.0
431.0             1  153293.000000  153293.0   153293.0  153293.0
655.0             1   12899.000000   12899.0    12899.0   12899.0

[119 rows x 5 columns]
```

<!-- grade_id: cell-a91dcfef5b191c30, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
assert track_stats.shape[0] == 119
assert track_stats.shape[1] == 5
assert track_stats.loc[1.0]['count'] == 45300
### END HIDDEN TEST
```

### Q11
>- From the total spotify data, make a dataframe called `wtf` based on the following rules:
>- Has records with `acousticness` >= .9 and `energy` >= 1
>- OR 
>- Has records with `danceability` <= 0 and `loudness` >= -3


<!-- grade_id: cell-3dc398f87e664a84, points: 0, grade: false, solution: true -->
```python
wtf = pd.DataFrame(spotify[spotify['acousticness'] >= .9])
wtf
```

Output:
```
acousticness                album_id  \
5              0.971  7CCwkPweMxKq8yWkVerH6T   
54             0.920  6hkfojKabPHetxgdYAAfWP   
73             0.922  5bFHtYB0sIpr08zR4TP9A5   
101            0.922  4LWaaqC9aAr5OhNsWqOZ3J   
151            0.916  7FEf6z8xu33s2hVp7sK0gv   
...              ...                     ...   
101895         0.973  1c5EcxSFk9Fok943GTAwFn   
101896         0.990  3Zs6jLiLhNzQcwJFHq677C   
101898         0.937  0EeoPgnFgCP1TNJsBdelJ7   
101899         0.906  119PlPHpuUWF8DYrQrjVHe   
101900         0.920  6LaSiuWQ4coE2tdhgeIJxx   

                                             analysis_url  \
5       https://api.spotify.com/v1/audio-analysis/4PrA...   
54      https://api.spotify.com/v1/audio-analysis/5J0a...   
73      https://api.spotify.com/v1/audio-analysis/2iE7...   
101     https://api.spotify.com/v1/audio-analysis/5dLo...   
151     https://api.spotify.com/v1/audio-analysis/0ZO4...   
...                                                   ...   
1018
```

<!-- grade_id: cell-ac893d61d314e8ee, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
assert wtf.shape[0] == 2
assert wtf.loc[29344]['album_id'] == '4mFFAzJgZ9886aTrcJC62R'

### END HIDDEN TEST
```